```python
# Initialize Otter
import otter
grader = otter.Notebook("roboworld-example.ipynb")
```

Für dieses Notebook benötigen Sie zwei Module ``roboworld`` und ``otter-grader``. Installieren Sie diese z.B. durch:

```
pip install otter-grader
pip install roboworld
```

# Roboterwelt

Roboter dienen uns heute in vielen Bereichen als Helfer und Assistenten.
Sie fertigen Autos und andere große Maschinen an, spielen gegeneinander Fußball und interagieren mit uns Menschen auf direktem Weg.
Roboter können dort eingesetzt werden, wo es für den Menschen zu gefährlich wird.
Zum Beispiel, hatte man versucht durch einen Roboter die Lage am Kernreaktor von Chernobyl besser zu analysieren.
Ein weiteres Beispiel ist die Entschärfung von Sprengstoff oder das Auffinden von Überlebenden, welche unter Gebäuden begraben wurden.

Wir werden heute einen Roboter durch eine zweidimensionale Welt bewegen.

Um die Welt und ihren Roboter anzuzeigen benötigen Sie das Modul ``roboworld``.

```python
import roboworld as rw
```

Dieser Befehl sorgt dafür, dass das *Paket* ``roboworld`` in unser Notebook importiert (eingefügt) wird. Um nicht immer den langen Namen ``roboworld`` tippen zu müssen, wählen wir die Abkürzung ``rw``. 

Lassen Sie uns eine Roboterwelt mit 5 Zeilen ``nrows`` und 9 Spalten ``ncols`` erzeugen:

```python
world = rw.new_world(nrows = 5, ncols = 9)
```

Wir sagen dem ``Modul``, dass es uns eine neue Welt erzeugen soll. Jetzt können wir dieser Welt ``world`` sagen, dass sie sich selbst bitte anzeigen soll:

```python
world.show()
```

***
***Aufgabe 5.*** 
Bewegen Sie Ihren Roboter ``robo`` auf das Ziel. Gehen Sie davon aus, dass der Roboter bzw. Sie wissen wo das Ziel ist. Nach Ihr Erstellung der Welt ist der Roboter nach **Norden (N)** ausgerichtet. Mir ``help(robo)`` können Sie erfragen, welche Operationen Ihr Roboter beherrscht.

```python
mycorridor = rw.new_world(nrows = 1, ncols = 10, robo_position=[0,0], goal_position=[0, 5])
robo = mycorridor.get_robo()
mycorridor.show()
```

```python
...
```

```python
grader.check("q5")
```

***

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

These are some submission instructions.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q5 (points: None)
```
>>> mycorridor.is_successful() == True
True
```
